```python
# Initialize Otter
import otter
grader = otter.Notebook("CT-04-Hurricanes.ipynb")
```

# 1 Hurrikans und Wörterbücher

## 1.1 Structure of Arrays vs Array of Structures (Zusatzinformation)

Zwei der Datenhaltungen die wir häufig antreffen sind sog. *Structure of Arrays SoA* und *Array of Structure* (siehe auch [AoS and SoA](https://en.wikipedia.org/wiki/AoS_and_SoA)).

### 1.1.1 Structure of Arrays
Bei *Structure of Arrays* (Struktur von Listen) wird je ein Typ von Datum für alle Objkte in einer Liste gehalten.
Alle Objekte zusammengenommen bilden sich dann aus den ganzen Listen, wobei jede Liste soviele Einträge besitzt, wie es Objekte gibt. Zum Beispiel:

```python
x = [1.2, 2.3, 4.5, ...]
y = [3.4, 2.1, 3.3, ...]
name = ['rect1', 'rect2', 'rect3', ...]
width = [1.4, 2.1, 3.3, ...]
height = [91.4, 28.1, 6.5, ...]
rectangles = {'x': x, 'y': y, 'name': name, 'width': width, 'height': height}
```

*Structure of Arrays SoA* werden dann verwendet, wenn man eine bestimmte Operation parallel auf allen Daten eines Typs ausführen möchte.
So verwenden wir auf Grafikkarten *Structures of Arrays SoA*.
Auch verwenden wir beim wissenschaftlichen Rechnen häufig *Structures of Arrays SoA*.


### 1.1.2 Array of Structures
Die zweite Form nennt man *Array of Structures AoS* (Liste von Strukturen).
Hierbei werden die Daten eines Objekts gebündelt und eine Liste mit all diesen Bündeln erstellt.
Zum Beispiel:

```python
rect1 = {'name': 'rect1', 'x': 1.2, 'y': 3.4, 'width': 1.4, 'height': 91.4}
rect2 = {'name': 'rect2', 'x': 2.3, 'y': 2.1, 'width': 2.1, 'height': 28.1}
rect3 = {'name': 'rect3', 'x': 4.5, 'y': 3.3, 'width': 3.3, 'height': 6.5}
...
objects = [rect1, rect2, rect3, ...]
```

Ein *Array of Structures* verwenden wir hingegen in der herkömmlichen sog. objektorientierten Programmierung.

Beide Formen haben Vor- und Nachteile.
In der folgenden Aufgabe werden wir sowohl mit *Structures of Arrays SoA* als auch mit einem *Array of Structures* arbeiten.

## 1.2 Aufgabenbeschreibung

Hurrikans, auch bekannt als Wirbelstürme, Zyklone oder Taifune, sind eine der mächtigsten Naturkräfte der Erde.
Aufgrund des vom Menschen verursachten Klimawandels haben Anzahl und Intensität von Hurrikanen zugenommen, was eine bessere Vorbereitung der vielen von ihnen verwüsteten Gemeinschaften erfordert.
Als besorgte Umweltschützer\*innen möchten Sie sich Daten über die stärksten Wirbelstürme ansehen, die sich ereignet haben.

Im folgenden sind die Daten zu den 34 stärksten atlantischen Hurrikans in einem *Structure of Arrays*  ``hurricanes`` in der Form eines Wörterbuchs ``dict`` bestehend aus Listen ``list`` zusammengefasst.
Die Daten umfassen:

* `name`: Namen der Hurrikans
* `month`: Monate, in denen die Wirbelstürme aufgetreten sind
* `year`: Jahre, in denen die Wirbelstürme aufgetreten sind
* `speed`: maximal anhaltende Winde (Kilometer pro Stunde) der Hurrikans
* `areas`: Liste der verschiedenen Gebiete, die von jedem der Hurrikans betroffen sind
* `deaths`: Gesamtzahl der durch jeden der Hurrikans verursachten Todesfälle

Die Daten sind so organisiert, dass die Daten für Index $i$ zum $i$-ten Hurrikan gehören. Anders ausgedrückt diese Daten gehören zu einem modellierten Objekt, den einen Hurrikan.

```python
# name of the hurricane
name = ['Cuba I', 'San Felipe II Okeechobee', 'Bahamas', 'Cuba II', 'CubaBrownsville', 'Tampico', 'Labor Day', 'New England', 'Carol', 'Janet', 'Carla', 'Hattie', 'Beulah', 'Camille', 'Edith', 'Anita', 'David', 'Allen', 'Gilbert', 'Hugo', 'Andrew', 'Mitch', 'Isabel', 'Ivan', 'Emily', 'Katrina', 'Rita', 'Wilma', 'Dean', 'Felix', 'Matthew', 'Irma', 'Maria', 'Michael']

# month of appearence
month = ['October', 'September', 'September', 'November', 'August', 'September', 'September', 'September', 'September', 'September', 'September', 'October', 'September', 'August', 'September', 'September', 'August', 'August', 'September', 'September', 'August', 'October', 'September', 'September', 'July', 'August', 'September', 'October', 'August', 'September', 'October', 'September', 'September', 'October']

# year of appearence
year = [1924, 1928, 1932, 1932, 1933, 1933, 1935, 1938, 1953, 1955, 1961, 1961, 1967, 1969, 1971, 1977, 1979, 1980, 1988, 1989, 1992, 1998, 2003, 2004, 2005, 2005, 2005, 2005, 2007, 2007, 2016, 2017, 2017, 2018]

# maximal wind speed
speed = [265, 260, 260, 275, 260, 260, 285, 260, 260, 275, 275, 260, 260, 275, 260, 275, 275, 290, 285, 260, 275, 280, 265, 265, 260, 275, 280, 285, 275, 275, 265, 280, 275, 260]

# regions the hurricane hit
areas = [['Central America', 'Mexico', 'Cuba', 'Florida', 'The Bahamas'], ['Lesser Antilles', 'The Bahamas', 'United States East Coast', 'Atlantic Canada'], ['The Bahamas', 'Northeastern United States'], ['Lesser Antilles', 'Jamaica', 'Cayman Islands', 'Cuba', 'The Bahamas', 'Bermuda'], ['The Bahamas', 'Cuba', 'Florida', 'Texas', 'Tamaulipas'], ['Jamaica', 'Yucatn Peninsula'], ['The Bahamas', 'Florida', 'Georgia', 'The Carolinas', 'Virginia'], ['Southeastern United States', 'Northeastern United States', 'Southwestern Quebec'], ['Bermuda', 'New England', 'Atlantic Canada'], ['Lesser Antilles', 'Central America'], ['Texas', 'Louisiana', 'Midwestern United States'], ['Central America'], ['The Caribbean', 'Mexico', 'Texas'], ['Cuba', 'United States Gulf Coast'], ['The Caribbean', 'Central America', 'Mexico', 'United States Gulf Coast'], ['Mexico'], ['The Caribbean', 'United States East Coast'], ['The Caribbean', 'Yucatn Peninsula', 'Mexico', 'South Texas'], ['Jamaica', 'Venezuela', 'Central America', 'Hispaniola', 'Mexico'], ['The Caribbean', 'United States East Coast'], ['The Bahamas', 'Florida', 'United States Gulf Coast'], ['Central America', 'Yucatn Peninsula', 'South Florida'], ['Greater Antilles', 'Bahamas', 'Eastern United States', 'Ontario'], ['The Caribbean', 'Venezuela', 'United States Gulf Coast'], ['Windward Islands', 'Jamaica', 'Mexico', 'Texas'], ['Bahamas', 'United States Gulf Coast'], ['Cuba', 'United States Gulf Coast'], ['Greater Antilles', 'Central America', 'Florida'], ['The Caribbean', 'Central America'], ['Nicaragua', 'Honduras'], ['Antilles', 'Venezuela', 'Colombia', 'United States East Coast', 'Atlantic Canada'], ['Cape Verde', 'The Caribbean', 'British Virgin Islands', 'U.S. Virgin Islands', 'Cuba', 'Florida'], ['Lesser Antilles', 'Virgin Islands', 'Puerto Rico', 'Dominican Republic', 'Turks and Caicos Islands'], ['Central America', 'United States Gulf Coast (especially Florida Panhandle)']]

# damage in million
damage = ['Schaden nicht aufgezeichnet', '100M', 'Schaden nicht aufgezeichnet', '40M', '27.9M', '5M', 'Schaden nicht aufgezeichnet', '306M', '2M', '65.8M', '326M', '60.3M', '208M', '1.42B', '25.4M', 'Schaden nicht aufgezeichnet', '1.54B', '1.24B', '7.1B', '10B', '26.5B', '6.2B', '5.37B', '23.3B', '1.01B', '125B', '12B', '29.4B', '1.76B', '720M', '15.1B', '64.8B', '91.6B', '25.1B']

# caused deaths
deaths = [90,4000,16,3103,179,184,408,682,5,1023,43,319,688,259,37,11,2068,269,318,107,65,19325,51,124,17,1836,125,87,45,133,603,138,3057,74]

hurricanes = dict(
    name=name, 
    month=month, 
    year=year, 
    speed=speed, 
    areas=areas, 
    damage=damage, 
    deaths=deaths)
#hurricanes
```

***
***Aufgabe 1 (Typumwandlung).*** Die Schäden wurden als Zeichenketten der Form `'100M'` bzw. `'20.9B'` abgelegt wobei `'M'` für Millionen und `'B'` für Billionen steht.
Damit lässt sich schlecht rechnen.

Schreiben Sie eine Funktion ``damage_to_float(damage)``, welche aus der Schadensliste ``damage`` bzw. ``hurricanes['damage']`` eine neue neue Liste aus Fließkommazahlen erzeugt und zurückgibt.
Ersetzen Sie dabei ``'Schaden nicht aufgezeichnet'`` durch ``None``.

Führen Sie Ihre Umwandlung duch und legen Sie die Liste in ``hurricanes['damage']`` ab.

```python
# BEGIN SOLUTION
def to_float(text):

    if text == 'Schaden nicht aufgezeichnet':
        return None
    elif text[-1] == 'M':
        return float(text[:-1:]) * 1e6
    else:
        return float(text[:-1:]) * 1e9
    
def damage_to_float(damage):
    return list(map(lambda text: to_float(text), damage))

hurricanes['damage'] = damage_to_float(hurricanes['damage'])
# END SOLUTION
```

```python
grader.check("q1")
```

***Aufgabe 2 (Kostenschätzung).*** Schreiben Sie eine Funktion ``estimate_costs(damage_list)``, die die Gesamtkosten, aller aufgezeichneten Hurrikans schätzt.
Fehlt eine Aufzeichnung werten Sie dies als 0 Kosten (Unterschätzung).
Führen Sie die Schätzung durch.
Nutzten Sie die Ergebnisse aus Aufgabe 1 (d.h. ``damage_list`` ist eine Liste aus Fließkommazahlen und ``None``).

```python
def estimate_costs(damage_list):
    # BEGIN SOLUTION
    estimate = 0
    for damage in damage_list:
        estimate += damage if damage != None else 0
    return estimate
    # END SOLUTION
```

```python
grader.check("q2")
```

***Aufgabe 3 (Durchschnittsgeschwindigkeit).*** Schreiben Sie eine Funktion ``avg_speed(speed_list)``, die die durchschnittliche Maximalgeschwindigkeit der Hurrikans berechnet. Dabei ist ``speed_list`` eine Liste aus Zahalen wie ``hurricanes['speed']``.

```python
def avg_speed(speed_list):
    # BEGIN SOLUTION
    if len(speed_list) == 0:
        return 0
    return sum(speed_list) / len(speed_list)
    # END SOLUTION
```

```python
grader.check("q3")
```

***Aufgabe 4 (Array of Structures).*** Schreiben Sie eine Funktion ``extract(hurricanes, i)``, welche den ``i``-ten Hurrikan Ihres Wörterbuchs ``hurricanes`` in ein Wörterbuch der folgenden Form umwandelt:

```python
{'name': 'Cuba I',
 'month': 'October',
 'year': 1924,
 'speed': 265,
 'areas': ['Central America', 'Mexico', 'Cuba', 'Florida', 'The Bahamas'],
 'damage': None,
 'deaths': 90}, ...
``` 

umwandelt.
**Hinweis:** Die Schlüssel des Ergebnisses gleichen den Schlüsseln des Wörterbuchs ``hurricanes``!

```python
def extract(hurricanes, i):
    # BEGIN SOLUTION
    hurricane = {}
    for key in hurricanes.keys():        
        hurricane[key] = hurricanes[key][i]
    return hurricane
    # END SOLUTION
```

```python
grader.check("q4")
```

***Aufgabe 5 (Array of Structures).*** Schreiben Sie eine Funktion ``to_array_of_structures(hurricanes)``, die aus dem Wörterbuch ``hurricanes`` eine Liste aus Wörterbüchern erstellt und zurückliefert. (Verwenden Sie Ihre zuvor definierte Funktion ``extract(hurricanes, i)``.

```python
def to_array_of_structures(hurricanes):
    # BEGIN SOLUTION
    hurricane_list = []
    for i in range(len(hurricanes['damage'])):
        hurricane_list.append(extract(hurricanes, i))
    return hurricane_list
    # END SOLUTION
hurricane_list = to_array_of_structures(hurricanes)
```

***Aufgabe 6 (Gebiete).*** Schreiben Sie eine Funktion ``extract_areas(hurricane_list)`` die alle Gebiete in denen ein Hurrikan aufgetreten ist aus Ihrer Liste aus Hurrikanobjekten zusammenfasst und zurückliefert. Es sollen keine doppelten Einträge entstehen. Welche Datenstruktur eignet sich dafür?

```python
def extract_areas(hurricane_list):
    # BEGIN SOLUTION
    areas = set()

    for hurricane in hurricane_list:
        areas |= set(hurricane['areas'])
    return areas
    # END SOLUTION
areas = extract_areas(hurricane_list)
```

***Aufgabe 7 (Gebiete).*** Wir möchten nun anstatt Hurrikans unsere Gebiete besser analyiseren. Dazu haben wir eine Funktion ``extract_area_information(hurricane_list)`` geschrieben, die ein Wörterbuch aus Wörterbüchern aus der Liste der Hurikans erzeugt.

Beschreiben Sie in Ihren Worten **was** ``extract_areas(hurricane_list)`` erzeugt. Führen Sie die Funktion aus. Was können Sie aus den Einträgen des Wörterbuchs herauslesen?

```python
def extract_area_information(hurricane_list):
    areas = {}
    for hurricane in hurricane_list:
        for areaname in hurricane['areas']:
            if areaname not in areas:
                areas[areaname] = {}
            
            if 'years' not in areas[areaname]:
                areas[areaname]['years'] = []
            if 'avg_speed' not in areas[areaname]:
                areas[areaname]['avg_speed'] = 0
            if 'avg_damage' not in areas[areaname]:
                areas[areaname]['avg_damage'] = 0
            if 'avg_deaths' not in areas[areaname]:
                areas[areaname]['avg_deaths'] = 0
            
            nareas = len(hurricane['areas'])
            areas[areaname]['years'].append(hurricane['year'])
            areas[areaname]['avg_damage'] += hurricane['damage'] / nareas if hurricane['damage'] != None else 0
            areas[areaname]['avg_deaths'] += hurricane['deaths'] / nareas
            areas[areaname]['avg_speed'] += hurricane['speed']
        
    for areaname, area in areas.items():
        areas[areaname]['avg_speed'] /= len(areas[areaname]['years'])
    return areas
extract_area_information(hurricane_list)
```

Output:
```
{'Central America': {'years': [1924,
   1955,
   1961,
   1971,
   1988,
   1998,
   2005,
   2007,
   2018],
  'avg_speed': 271.6666666666667,
  'avg_damage': 26816216666.666668,
  'avg_deaths': 7451.516666666667},
 'Mexico': {'years': [1924, 1967, 1971, 1977, 1980, 1988, 2005],
  'avg_speed': 270.7142857142857,
  'avg_damage': 2058183333.3333333,
  'avg_deaths': 402.6833333333334},
 'Cuba': {'years': [1924, 1932, 1933, 1969, 2005, 2017],
  'avg_speed': 272.5,
  'avg_damage': 17522246666.666668,
  'avg_deaths': 785.9666666666666},
 'Florida': {'years': [1924, 1933, 1935, 1992, 2005, 2017],
  'avg_speed': 275.0,
  'avg_damage': 29438913333.333336,
  'avg_deaths': 209.06666666666663},
 'The Bahamas': {'years': [1924, 1928, 1932, 1932, 1933, 1935, 1992],
  'avg_speed': 268.57142857142856,
  'avg_damage': 8870580000.0,
  'avg_deaths': 1682.2333333333331},
 'Lesser Antilles': {'years': [1928, 1932, 1955, 2017],
  'avg_speed': 271.25,
  'avg_damage': 18384566666.666668,
  'avg_deaths': 2640
```

_Type your answer here, replacing this text._

Die Schlüssel des Wörterbuchs welches uns die Funktion ``extract_area_information(hurricane_list)`` erstellt sind die Namen der Gebiete. Jeder Wert ist ein Wörterbuch welches die Informationen über das Gebiet bzgl. Auswirkungen der Hurrikans enthält. Da aus den Daten nicht hervorgeht wie viel Schaden und Tote und für ein Gebiet insgesamt entstanden sind, wird dieser auf die Gebiete in gleichen Anteilen aufgeteilt (was natürlich eine grobe Schätzung ist). Zudem wird die durschnittliche Maximalgeschwindigkeit der Hurrikans, welche das Gebiet getroffen haben berechnet.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: None)
```
>>> damage_to_float(['Schaden nicht aufgezeichnet', '1M', '2.3B']) == [None, 1000000.0, 2300000000.0]
True
```
```
>>> damage_to_float(['Schaden nicht aufgezeichnet', '102.13M', '2.3B', 'Schaden nicht aufgezeichnet']) == [None, 102130000.0, 2300000000.0, None]
True
```


Tests for q2 (points: None)
```
>>> estimate_costs([None, 1.2, 3.5, None]) == 4.7
True
```
```
>>> estimate_costs([]) == 0.0
True
```
```
>>> estimate_costs([None]) == 0.0
True
```


Tests for q3 (points: None)
```
>>> avg_speed([]) == 0
True
```
```
>>> avg_speed([1.3, 1.4]) == 1.35
True
```


Tests for q4 (points: None)
```
>>> test_hurricanes = {'name': ['Cuba I',
...   'San Felipe II Okeechobee',
...   'Bahamas',
...   'Cuba II',
...   'CubaBrownsville',
...   'Tampico',
...   'Labor Day',
...   'New England',
...   'Carol',
...   'Janet',
...   'Carla',
...   'Hattie',
...   'Beulah',
...   'Camille',
...   'Edith',
...   'Anita',
...   'David',
...   'Allen',
...   'Gilbert',
...   'Hugo',
...   'Andrew',
...   'Mitch',
...   'Isabel',
...   'Ivan',
...   'Emily',
...   'Katrina',
...   'Rita',
...   'Wilma',
...   'Dean',
...   'Felix',
...   'Matthew',
...   'Irma',
...   'Maria',
...   'Michael'],
...  'month': ['October',
...   'September',
...   'September',
...   'November',
...   'August',
...   'September',
...   'September',
...   'September',
...   'September',
...   'September',
...   'September',
...   'October',
...   'September',
...   'August',
...   'September',
...   'September',
...   'August',
...   'August',
...   'September',
...   'September',
...   'August',
...   'October',
...   'September',
...   'September',
...   'July',
...   'August',
...   'September',
...   'October',
...   'August',
...   'September',
...   'October',
...   'September',
...   'September',
...   'October'],
...  'year': [1924,
...   1928,
...   1932,
...   1932,
...   1933,
...   1933,
...   1935,
...   1938,
...   1953,
...   1955,
...   1961,
...   1961,
...   1967,
...   1969,
...   1971,
...   1977,
...   1979,
...   1980,
...   1988,
...   1989,
...   1992,
...   1998,
...   2003,
...   2004,
...   2005,
...   2005,
...   2005,
...   2005,
...   2007,
...   2007,
...   2016,
...   2017,
...   2017,
...   2018],
...  'speed': [265,
...   260,
...   260,
...   275,
...   260,
...   260,
...   285,
...   260,
...   260,
...   275,
...   275,
...   260,
...   260,
...   275,
...   260,
...   275,
...   275,
...   290,
...   285,
...   260,
...   275,
...   280,
...   265,
...   265,
...   260,
...   275,
...   280,
...   285,
...   275,
...   275,
...   265,
...   280,
...   275,
...   260],
...  'areas': [['Central America', 'Mexico', 'Cuba', 'Florida', 'The Bahamas'],
...   ['Lesser Antilles',
...    'The Bahamas',
...    'United States East Coast',
...    'Atlantic Canada'],
...   ['The Bahamas', 'Northeastern United States'],
...   ['Lesser Antilles',
...    'Jamaica',
...    'Cayman Islands',
...    'Cuba',
...    'The Bahamas',
...    'Bermuda'],
...   ['The Bahamas', 'Cuba', 'Florida', 'Texas', 'Tamaulipas'],
...   ['Jamaica', 'Yucatn Peninsula'],
...   ['The Bahamas', 'Florida', 'Georgia', 'The Carolinas', 'Virginia'],
...   ['Southeastern United States',
...    'Northeastern United States',
...    'Southwestern Quebec'],
...   ['Bermuda', 'New England', 'Atlantic Canada'],
...   ['Lesser Antilles', 'Central America'],
...   ['Texas', 'Louisiana', 'Midwestern United States'],
...   ['Central America'],
...   ['The Caribbean', 'Mexico', 'Texas'],
...   ['Cuba', 'United States Gulf Coast'],
...   ['The Caribbean', 'Central America', 'Mexico', 'United States Gulf Coast'],
...   ['Mexico'],
...   ['The Caribbean', 'United States East Coast'],
...   ['The Caribbean', 'Yucatn Peninsula', 'Mexico', 'South Texas'],
...   ['Jamaica', 'Venezuela', 'Central America', 'Hispaniola', 'Mexico'],
...   ['The Caribbean', 'United States East Coast'],
...   ['The Bahamas', 'Florida', 'United States Gulf Coast'],
...   ['Central America', 'Yucatn Peninsula', 'South Florida'],
...   ['Greater Antilles', 'Bahamas', 'Eastern United States', 'Ontario'],
...   ['The Caribbean', 'Venezuela', 'United States Gulf Coast'],
...   ['Windward Islands', 'Jamaica', 'Mexico', 'Texas'],
...   ['Bahamas', 'United States Gulf Coast'],
...   ['Cuba', 'United States Gulf Coast'],
...   ['Greater Antilles', 'Central America', 'Florida'],
...   ['The Caribbean', 'Central America'],
...   ['Nicaragua', 'Honduras'],
...   ['Antilles',
...    'Venezuela',
...    'Colombia',
...    'United States East Coast',
...    'Atlantic Canada'],
...   ['Cape Verde',
...    'The Caribbean',
...    'British Virgin Islands',
...    'U.S. Virgin Islands',
...    'Cuba',
...    'Florida'],
...   ['Lesser Antilles',
...    'Virgin Islands',
...    'Puerto Rico',
...    'Dominican Republic',
...    'Turks and Caicos Islands'],
...   ['Central America',
...    'United States Gulf Coast (especially Florida Panhandle)']],
...  'damage': ['Schaden nicht aufgezeichnet',
...   '100M',
...   'Schaden nicht aufgezeichnet',
...   '40M',
...   '27.9M',
...   '5M',
...   'Schaden nicht aufgezeichnet',
...   '306M',
...   '2M',
...   '65.8M',
...   '326M',
...   '60.3M',
...   '208M',
...   '1.42B',
...   '25.4M',
...   'Schaden nicht aufgezeichnet',
...   '1.54B',
...   '1.24B',
...   '7.1B',
...   '10B',
...   '26.5B',
...   '6.2B',
...   '5.37B',
...   '23.3B',
...   '1.01B',
...   '125B',
...   '12B',
...   '29.4B',
...   '1.76B',
...   '720M',
...   '15.1B',
...   '64.8B',
...   '91.6B',
...   '25.1B'],
...  'deaths': [90,
...   4000,
...   16,
...   3103,
...   179,
...   184,
...   408,
...   682,
...   5,
...   1023,
...   43,
...   319,
...   688,
...   259,
...   37,
...   11,
...   2068,
...   269,
...   318,
...   107,
...   65,
...   19325,
...   51,
...   124,
...   17,
...   1836,
...   125,
...   87,
...   45,
...   133,
...   603,
...   138,
...   3057,
...   74]}
>>> 
>>> extract(test_hurricanes,3) == {'name': 'Cuba II',
...  'month': 'November',
...  'year': 1932,
...  'speed': 275,
...  'areas': ['Lesser Antilles',
...   'Jamaica',
...   'Cayman Islands',
...   'Cuba',
...   'The Bahamas',
...   'Bermuda'],
...  'damage': '40M',
...  'deaths': 3103}
True
```
