```python
# Initialize Otter
import otter
grader = otter.Notebook("CT-02-Pyramiden.ipynb")
```

# 1 Aufwärmen

***
***Aufgabe 1 (Listen).*** Erstellen Sie auf unterschiedliche Weise eine Liste ``small_numbers``, welche die Zahlen 0 bis 9 enthält (beides inklusive).

```python
...
```

```python
grader.check("q1")
```

***Aufgabe 2 (Listen).*** Erstellen Sie eine Liste ``pyramid_list`` der Länge 10, die an der ``i``-ten Position eine Liste der Zahlen ``[0,1,...,i]`` enthält. Das Ergebnis ist eine Liste aus Listen.

```python
pyramid_list = ...
...
pyramid_list
```

```python
grader.check("q2")
```

***Aufgabe 3 (Zeichenketten).*** Wir bezeichnen eine Zeichenfolge Palindrom wenn sie vorwärts und rückwärts gelesen das gleiche Wort ergibt. Z.B. ist ``anna`` oder auch ``aabaa`` ein Palindrom. Schreiben Sie einen Code der für eine beliebige Zeichenkette ``text`` prüft ob es sich um ein Palimdrom handelt. Testen Sie Ihren Code.

```python
text = 'anna'
...
```

***Aufgabe 4 (Funktionen).*** Schreiben Sie eine Funktion ``is_even(n)``, die Ihnen wahr ``True`` zurückgibt wenn ``n`` eine gerade Zahl ist und sonst ``False``. Gehen Sie davon aus, dass ``n`` eine ganze Zahl ist.

```python
def is_even(n):
    ...
```

```python
grader.check("q4")
```

***Aufgabe 5 (Funktionen, Kontrollstruturen).*** Schreiben Sie eine Funktion ``specify_number(n)``, die einen Text ausgibt, der besagt ob die ganze Zahl ``n`` gerade oder ungerade und positiv oder negativ ist (0 werten wir als positiv).

```python
def specify_number(n):
    ...
```

***Aufgabe 6. (Kontrollstrukturen)*** Geben Sie 20 Mal ``'Die richtige Antwort ist 42!'`` aus.

```python
...
```

***Aufgabe 7 (Listen, Kontrollstrukturen).*** Erstellen Sie eine Liste ``numbers``, die alle ungeraden positiven Zahlen, welche durch 7 teilbar und kleiner als 1000 sind.

```python
numbers = ...
...
```

```python
grader.check("q7")
```

***Aufgabe 8 (Funktionen, Kontrollstrukturen).*** Schreiben Sie eine Funktion ``sum_squares(n)``, die Ihnen die Summe aller Quadratzahlen von 1 bis einschließlich ``n`` berechnet und zurückgibt. Gehen Sie davon aus, dass ``n`` eine natürliche Zahl $\mathbb{N}$ (0 eingeschlossen) ist.

```python
def sum_squares(n):
    ...
```

```python
grader.check("q8")
```

***

# 2 Münzliebhaber\*in

Sie möchten heute gerne einen bestimmten Geldbetrag ``money`` in Münzen ``coins`` umwandeln.
Doch zuvor möchten Sie wissen wie viele Münzen Sie denn minimal benötigen, schließlich möchten Sie sich nicht zu Tode Schleppen.

Unsere Währung ist der Euro. Es gibt 2 und 1 Euro Münzen (200, 100 Cent) sowie 50, 20, 10, 5, 2 und 1 Cent Münzen. Sie brauchen die wenigsten Münzen für einen bestimmten Geldbetrag wenn Sie die Anzahl der wertvollsten Münzen maximieren. Warum?

***
***Aufgabe 9 (Funktionen, Kontrollstrukturen).*** Schreiben Sie eine Funktion ``count_coins(money)``, die Ihnen die Anzahl der Münzen, die Sie **minimal** benötigen, berechnet! Dabei soll ``money`` der Geldbetrag in Euro sein, welchen Sie in Münzen umwandeln möchten. Gehen Sie davon aus, dass ``money`` eine positive Fließkommazahl ``float`` ist und runden Sie (mit ``round()``) auf den nähersten Geldbetrag, z.B. ``100.2356`` wird auf ``100.24`` gerundet und ``50.224`` auf ``50.22``.

```python
def count_coins(money):
    ...
```

```python
grader.check("q9")
```

***

# 3 Pyramidenbau (schwer)

## 3.1 Einleitung

Sie kennen vielleicht noch die alten sogenannten Jump & Run Konsolenspiele wie etwa Mario, Sonic und viele mehr.

Das Spielprinzip ist einfach: In einer zweidimensionalen Welt, in der sich die Spielfigur nur nach vorne bzw. zurück und hoch bzw. runter bewegen kann, muss diese Figur ein bestimmtes (örtliches) Ziel erreichen. Auf dem Weg lauern Gefahren und Ihre Figur muss Hindernisse überspringen. Sie durchlaufen mehrere Levels, deren Schwierigkeitsgrad ansteigt.

Interessanterweise gibt es diese Spiele heute immernoch - sie haben einen gewissen Kultstatus. Besonders begehrt sind Spiele bei denen Spieler\*innen ihre eigenen Welten erstellen können und diese dann von anderen Spieler\*innen gespielt und bewertet werden.

Ist die Spielwelt aus einfachen Elementen welche in einem regulären Gitter angeordnet sind aufgebaut, so kann eine solche Welt als lesbarer Text in eine Datei geschrieben werden. Ein Beispiel: Stellen Sie sich vor, folgender Text stünde in einer Textdatei:

```txt
     #  #
    ##  ##
 P ###  ###   *
######  #### ###
```

Dabei könnten wir definieren, dass ``P`` die Startposition der Spielfigur ist, ``#`` quadratische Bauklötze und ``*`` das zu erreichende Ziel ist. Diese Information könnte von einem Programm in eine entsprechende Welt umgewandelt werden.

![Jump & Run Gamescreenshot](figs/jump-and-run-screenshot.png)

Auch hierbei handelt es sich um eine **Interpretation**, eine ganze Spielewelt durch eine Folge von Zeichen **repräsentiert**! 

## 3.2 Aufgabenbeschreibung

Ihr Aufgabe ist es ein ``Python``-Programm zu schreiben, welches Ihnen eine Pyramide der Form

```txt
   #  #
  ##  ##
 ###  ###
####  ####
```

erzeugt, wobei die Höhe der Pyramide als Eingabeargument übergeben werden und zwischen 1 und 8 liegen soll.


***
***Aufgabe 10 (Funktionen, Kontrollstrukturen).*** Entwickeln Sie eine Funktion ``pyramid(height)`` die die oben beschriebene Pyramide der Höhe ``height`` als Zeichenkette ``str`` erzeugt und zurückgibt. Gehen Sie davon aus, dass ``height`` eine nicht zu große natürliche Zahl (0 ausgeschlossen) ist.

Zeichenketten können Sie mit dem ``+``-Operator verketten. Eine neue Zeile können Sie mit der Zeichenkette ``'\n'`` generieren. Zum Beispiel:

```python
print('abcd' + 'efg' + '\n' + 'hijk')
```

Oft wollen wir überflüssig Zeilen am Ende einer Zeichenkette (oder auch Textdatei) vermeiden.

```python
print('abcd' + 'efg' + '\n' + 'hijk\n\n\n')
```

```python
def pyramid(height):
    ...
```

```python
grader.check("q10")
```

***

***
***Aufgabe 11 (Funktionen, Kontrollstrukturen).*** Entwickeln Sie eine Funktion ``print_pyramid()`` die nach einer Eingabe (der Höhe) frägt und Ihnen die oben beschriebene Pyramide der Höhe dieser Eingabe erzeugt genauer gesagt als ausgibt. Verwenden Sie Ihre Funktion ``pyramid(height)``. Bei fehlerhafter Eingabe sollte stattdessen eine Fehlermeldung ausgegeben werden. Probieren Sie Ihre Funktion aus.

Um einen Wert in einem Notebook einzulesen, können Sie die *built-in*-Funktion ``input()`` verwenden, zum Beispiel:

```python
height = input('height = ')
print('data type:' + str(type(height)))
print('value: ' + str(height))
```

Mit der folgenden Funktion können Sie testen ob eine Zeichenkette tatsächlich eine natürliche Zahl $\mathbb{N}$ (0 ausgeschlossen) ist.

```python
def is_natural_number(value):
    try:
        int(value)
        return True
    except ValueError:
        return int(value) > 0
```

```python
def print_pyramid():
    ...
```

***

```python
print_pyramid()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: None)
```
>>> small_numbers == [0,1,2,3,4,5,6,7,8,9]
True
```


Tests for q10 (points: None)
```
>>> pyramid(1) == '#  #'
True
```
```
>>> pyramid(2) == ' #  #\n##  ##' or pyramid(2) ==  ' #  # \n##  ##'
True
```
```
>>> pyramid(3) == '  #  #\n ##  ##\n###  ###' or pyramid(3) ==  '  #  #  \n ##  ## \n###  ###'
True
```
```
>>> pyramid(4) == '   #  #\n  ##  ##\n ###  ###\n####  ####' or pyramid(4) ==  '   #  #   \n  ##  ##  \n ###  ### \n####  ####'
True
```


Tests for q2 (points: None)
```
>>> pyramid_list == [list(range(i+1)) for i in range(10)]
True
```


Tests for q4 (points: None)
```
>>> is_even(2) == True
True
```
```
>>> is_even(1202) == True
True
```
```
>>> is_even(-1202) == True
True
```
```
>>> is_even(-3) == False
True
```
```
>>> is_even(-137) == False
True
```


Tests for q7 (points: None)
```
>>> numbers == [i for i in range(1000) if i % 2 == 1 and i % 7 == 0]
True
```


Tests for q8 (points: None)
```
>>> sum_squares(1) == 1
True
```
```
>>> sum_squares(5) == 5**2 + 4**2 + 3**2 + 2**2 + 1**2
True
```
```
>>> sum_squares(9) == 9**2 + 8**2 + 7**2 + 6**2 + 5**2 + 4**2 + 3**2 + 2**2 + 1**2
True
```
```
>>> sum_squares(0) == 0
True
```


Tests for q9 (points: None)
```
>>> count_coins(0.13) == 3
True
```
```
>>> count_coins(100) == 50
True
```
```
>>> count_coins(100.31) == 53
True
```
```
>>> count_coins(100.316) == count_coins(100.32)
True
```
```
>>> count_coins(100.314) == count_coins(100.31)
True
```
