```python
# Initialize Otter
import otter
grader = otter.Notebook("CT-06-Uebungen_Template.ipynb")
```

# Vorlesung "Computational Thinking"
### Woche 6 - Übungen
#### Prof. Dr.-Ing. Martin Hobelsberger

### Lernziele dieser Einheit

* Wiederholungen zu Python Grundlagen

## Übungen zum Warm-Up

***Aufgabe 1.*** Schreiben Sie eine Funktion ``index_at(text, char)``, welche die erste vorkommende Position eines gesuchten Zeichens ``char`` in einer Zeichenkette ``text`` zurückgibt. Wir beginnen bei 0 zu zählen, d.h. 

```python
index_at("abc", "a") == 0
```

```python
def index_at(text, char):
    ...
```

```python
grader.check("q1")
```

```python
text = 'May the force be with you!'
char = 'f'
index = index_at(text, char)
print(f'Der Buchstabe {zeichen} ist an Stelle {index} im Text: {text}')
```

***Aufgabe 2:*** Schreiben Sie eine Yoda-Funktion ``yoda(text)``. Diese gibt einen übergebenen Satz ``text`` mit umgekehrten Wörtern wieder zurück.

```python
yoda('Ich bin bereit') == 'bereit bin Ich'
```

**Tipp:** Sehen Sie sich die String-Methoden ``join()`` und ``split()`` einer Zeichenkette ``str`` genauer an. 

```python
def yoda(text):
    ...
```

```python
grader.check("q2")
```

```python
yoda('Hallo mein Name ist Hase')
```

***Aufgabe 3.*** Schreiben Sie ein Programm, das eine Liste (mit einer belibigen Anzahl an numerischen Elementen z.B.,

```python
numbers = [2, 42, 23.2, 34, 22332.5]
```

als Eingabe entgegen nimmt und alle Elemente der Liste mit einem **geraden Index** zu einem Summenwert sowie alle Elemente mit **ungeraden Index** zu einem Summenwert addiert. Es sollen beide Summenwerte ausgegeben werden.

```python
...
```

***Aufgabe 4.*** Schreiben Sie eine Funktion ``mask(numbers)``, die eine Liste (mit ganzen Zahlen als Elemente) entgegen nimmt, für jedes Element der Liste überprüft, ob es gerade oder ungerade ist und eine neue Liste mit den entsprechenden String-Elementen ``'gerade'`` bzw. ``'ungerade'`` zurückliefert.

Zum Beispiel sollte gelten:

```python
numbers = [7,34,2,22,31]
mask(numbers) == ['ungerade', 'gerade', 'gerade', 'gerade', 'ungerade']
```

```python
def mask(numbers):
    ...
```

```python
grader.check("q4")
```

```python
numbers = [7, 34, 2, 22, 31]
mask(numbers)
```

***Aufgabe 5.*** Schreiben Sie eine Funktion ``median(numbers)``, das eine Liste (mit einer belibigen Anzahl ``n`` an ganzzahligen Elementen; z.B. ``numbers = [2, 42, 232, 34, 5]``) als Eingabe entgegen nimmt, den [Median](https://de.wikipedia.org/wiki/Median) der Listenelemente berechnet und ausgibt.

**Tipp:** Mit der Methode ``sort`` einer Liste ``list`` könen Sie eine entsprechende Liste aus Zahlen sortieren.

```python
def median(numbers):
    ...
```

```python
grader.check("q5")
```

***Aufgabe 6.*** Schreiben Sie eine Funktion die eine Liste (mit ganzen Zahlen als Elemente) entgegen nimmt und darüber die Quersumme bildet. Die Funktion soll auf der Konsole den Wert der Quersumme, die Anzahl der Elemente und ob die Quersumme gerade oder ungerade ist ausgeben. Als Rückgabewert wird die Quersumme geliefert. 

```python
...
```

***Aufgabe 7. (Persönliche Daten aus Zeichenketten ausgeben)*** In einer Textdatei sind Zeile für Zeile Angaben zu einer Person gespeichert:

`'Name: Dörte Weber;'`

`'Adresse: Kobelweg 12, 86391 Stadtbergen;'`

`'Alter: 66;'`

Die Zeichenketten dieser Datei werden zeilenweise eingelesen in die Variable `row`. Schreiben Sie einen Code, der mit `print` nur die persönlichen Daten ausgibt (d.h. den String hinter dem Doppelpunkt und ohne Leerzeichen beginnend).

Hier im Beispiel, muss also aus `row` der gewünschte String extrahiert werden: 
```'Dörte Weber;'```. Teste dein Programm auch für die anderen Zeilen.

**Tipp:** Eventuell hilft Ihnen ``index_at`` von Aufgabe 1 und die Methode ``strip`` der Klasse ``str``.

```python
row = "Name: Dörte Weber;"
```

```python
...
```

***Aufgabe 8. (Kürzen von Zeichenketten)*** Für einen Fashion-Blog soll eine kurze Vorschau der Artikel im Umfang von 35 Zeichen angezeigt werden.
Der Text soll dabei mit drei Punkten `...` enden, um anzuzeigen, dass es noch mehr Inhalt gibt.

Schreiben Sie den Code, um einen `post` wie beschrieben neu zu formatieren und auszugeben. Für

```python
post = "Die Herbst- und Winter-Saison ist eingeleitet und die Duftkerzen sind bereits angezündet. Jetzt fehlen also nur noch ein paar Parfums, die perfekt in die kältere Jahreszeit passen (aber in meinem Fall eigentlich auch das ganze Jahr über) – hier sind also meine aktuellen Lieblinge!"
```

Ist die gewünschte Ausgabe gleich

```python
Die Herbst- und Winter-Saison ist e...
```

```python
length = len(post)
n = 35
if length < n:
    print(post)
else:
    print(f'{post[:n]}...')
```

***Aufgabe 9. (Verbessertes Kürzen von Zeichenketten)*** Können Sie Aufgabe 8 so lösen, dass die drei Punkte nach einem Leerzeichen eingefügt werden, um zu verhindern, dass Texte mitten in einem Wort abgeschnitten werden?

Ausgabe: 
```Die Herbst- und Winter-Saison ist...```

```python
...
```

```python
"abc".find('b')
```

***Aufgabe 10. (Glückspiel)*** In ``w`` sei der Wert eines Würfelwerts gespeichert. Der Würfel kann die Werte 1, 2, 3, 4, 5 und 6 annehmen - Zum Beispiel gelte ``w==4``.

Schreibe Programmcode, der mit Hilfe des Modulo-Operators prüft, ob der Wert von ``w`` gerade oder ungerade ist. Das Programm soll z.B. ```3 ist ungerade``` oder ```4 ist gerade``` ausgeben.

**Beispielausgabe:**

```4 ist gerade```

```python
w = 4
```

```python
result = 'gerade' if w % 2 == 0 else 'ungerade'
print(f'{w} ist {result}')
```

***Aufgabe 11. (Glückspiel 2)*** Nun wir der Würfen zweimal hintereinander geworfen. Wenn beide Würfelwerte ungerade sind, erhält der Spieler eine Gewinnsumme von 10 Euro. Wenn nur ein Würfelwert ungerade ist, beträgt die Gewinnsumme 0 Euro. Wenn beide Würfelwerte gerade sind, macht der Spieler einen Verlust von 10 Euro.

Ihre Aufgabe ist es, ein Programm zu schreiben, welches die Gewinnsumme bzw. den Verlust ausgibt.

**Beispielausgabe für w1=1 und w2=3:**
```Gewinnsumme beträgt 10 Euro.```

**Beispielausgabe für w1=1 und w2=4:**
```Gewinnsumme beträgt 0 Euro.``` 

**Beispielausgabe für w1=2 und w2=4:**
```Verlust beträgt 10 Euro.``` 

Testen Sie Ihren Programmcode für die verschiedenen Würfelwerte.

```python
w1 = 1
w2 = 3
```

```python
...
```

***Aufgabe 12. (Versandkosten)*** Für einen Food-Online-Shop sollen Sie ein Programm schreiben, dass die Versandkosten berechnet. Die Variable ``bill`` enthält die Summe der Preise aller gekauften Waren. Die Versandkosten betragen ``8.95`` Euro. Ab einem Einkaufswert von ``40`` Euro betragen die Versandkosten nur ``5.95`` Euro; ab einem Einkaufswert von ``50`` Euro werden keine Versandkosten berechnet.

Deine Aufgabe ist es, ein Programm zu schreiben, dass die Versandkosten zum Preis (also zu der Zwischensumme) addiert und den (neuen) Gesamtpreis ausgibt. Weiterhin soll das Programm ggf. ausgeben, für wie viel Geld noch Waren gekauft werden müssen, damit die Bestellung versandkostenfrei geliefert wird bzw. nur 5,95 Euro beträgt.

**Hinweis**: Nutze die Funktion `round()`. Der zweite Parameter für die Funktion `round()` gibt an, auf wie viele Nachkommastellen gerundet werden soll. Bsp: `round(1.2934,2)` ergibt `1.29`. Im Zusammenhang mit der Ausgabe von Geldbeträgen kann diese Funktion sehr sinnvoll sein.

**Beispielausgabe für die Zwischensumme von 47.60 Euro:**

``Wenn Sie weitere Waren im Wert von 2.4 Euro einkaufen, erfolgt der Versand kostenfrei.
Gesamtpreis: 53.55 Euro``

**Beispielausgabe für die Zwischensumme von 36.25 Euro:**

``Wenn Sie weitere Waren im Wert von 3.75 Euro einkaufst, zahlen Sie für die Versandkosten nur 5.95 Euro. Wenn Sie weitere Waren im Wert von 13.75 Euro einkaufst, erfolgt der Versand kostenfrei.
Gesamtpreis: 45.2 Euro``

```python
def print_shipping_message(bill):
    ...
```

```python
bill = 47.60
print_shipping_message(bill)
```

```python
bill = 36.25
print_shipping_message(bill)
```

***Aufgabe 13. (Öffnungszeiten)*** Schreiben Sie eine Funktion ``greetings(sunday, holiday)``, die Kunden eines Geschäfts grüßt.

1. Handelt es sich beim aktuellen Tag um einen Sonntag, d.h. ``sunday==True`` so lautet die Ausgabe:
``Wir haben geschlossen. Einen schönen Sonntag.``.
2. Sofern es sich nur um einen Feiertag handelt, d.h. ``holiday==True`` und ``sunday==False`` so lautet die Ausgabe:
``Wir haben geschlossen.``.
3. Sonst lautet die Ausgabe:
``Willkommen!``.


**Ausgabe für Sonntag:**
```Wir haben geschlossen. Einen schönen Sonntag.```

**Ausgabe für einen Feiertag:**
```Wir haben geschlossen.```

**Ausgbe für einen Werktag:**
```Willkommen!```


Mache dir Gedanken dazu, inwiefern die Reihenfolge der Bedingungen eine Rolle spielt!

```python
def greetings(sunday, holiday):
    text = ''
    ...
    print(text)  
```

```python
greetings(True, True)
greetings(True, False)
greetings(False, True)
greetings(True, False)
greetings(False, False)
```

***Aufgabe 14. (Häufigkeitsanalyse)*** Schreiben Sie eine Funktion ``count_words(text, pattern)``, die die Anzahl der Vorkommnisse der Zeichenfolge ``pattern`` im Text ``text`` berechnet und zurückliefert. Dabei soll die Groß- und Kleinschreibweise ignoriert werden. Hierfür kannst du die `.lower()`-Funktion verwenden.

Öffne anschließend den text von TopicModellingPaper18, und ermittle, wie oft die Wörter "topic", "word" und "frequency" in diesem Paper vorkommen.

**Hinweis:** Die Referenzen, Angaben zu den Autoren, Überschriften zählen wir jetzt mit dazu. In bestimmten Analyseverfahren werden diese Textelemente häufig ausgeschlossen.

**Hinweis:** Mit `.count('data')` wird ausgegeben, wie oft der Teilstring `data` in einem String vorkommt. Wichtig: Diese Funktion unterscheidet zwischen Groß- und Kleinschreibung:

```python
def count_words (text, pattern):
    ...
    
with open("TopicModellingPaper18.txt", encoding = "utf-8") as file:
    content = file.read()
    ...
```

```python
#Beispiel mit .count()
text = "Data data-science dataanalyst data"
print(text.count("data"))
```

Die Ausgabe Ihrer Analyse sollte so aussehen:

``190-mal kommt 'topic' vor.
  91-mal kommt 'word' vor.
  2-mal kommt 'frequency' vor.``

**Hinweis:** Falls noch nicht gemacht, nutze die `for-Schleife`, um die einzelnen Wörter zu analysieren. 

```python

```

***Weitere Möglichkeiten:***
Ausgabe der Analyse in eine eigene Datei schreiben oder die Liste an Wörtern, die analysiert werden sollen, beliebig erweitern lassen, indem Wörter individuell eingegeben werden können. Experimentieren Sie selbst ein wenig herum!

```python

```

***Aufgabe 15. (Häufigkeitsanalyse aller Wörter)*** Schreiben Sie eine Funktion ``count_all_words(text)``, welche ein Wörterbuch ``dict`` zurückliefert wobei ``key`` ein Wort in ``text`` ist und dessen Wert ``value`` die Anzahl der Vorkommnisse in ``text`` ist. Verwenden Sie erneut die Datei "TopicModellingPaper18.txt" zum Testen.

**Beachte:** Der Text enthält natürlich viele Satzzeichen. Sie müssen daher den Text zuerst an einem Leerzeichen trennen, und anschließend die ``word_cleaner`` Funktion (die bereits vorgegeben ist) verwenden, um Satzzeichen aus einem Wort zu entfernen. Diese Funktion entfernt alle Satzzeichen aus einem Wort, und gibt das Wort in Kleinbuchstaben zurück. Versuchen Sie als Erstes, diese Funktion nachzuvollziehen!
Schreiben Sie daraufhin Ihre Funktion ``count_all_words(text)``. Wir ignorieren alle Wörter die Sonderzeichen enthalten, z.B. ``'wouldn't'`` oder ``'http:\\www.hm.cs.edu'``.

**Hinweis 1:** Die Referenzen, Angaben zu den Autoren, Überschriften zählen wir mit dazu. 

**Hinweis 2:** Beachte, dass sich der leere String in deine Wörter einschleichen kann. Dieser zählt nicht als Wort!

```python
def count_words(text, pattern):
    return text.lower().count(pattern.lower())

def word_cleaner(word):
    allowed_characters = "abcdefghijklmnopqrstuvwxyzüöäß"
    w = ""
    for character in word.lower():
        if character in allowed_characters:
            w += character
    return w
        
def count_all_words(text):
    ...
    
with open("TopicModellingPaper18.txt", encoding="utf-8") as file:
    content = file.read()
    ...
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: None)
```
>>> index_at("Hallo mein Name ist Hase", "a") == 1
True
```
```
>>> index_at("Hallo mein Name ist Hase", "m") == 6
True
```
```
>>> index_at("Hallo mein Name ist Hase", "H") == 0
True
```
```
>>> index_at("Hallo mein Name ist Hase", "s") == 17
True
```


Tests for q2 (points: None)
```
>>> yoda('') == ''
True
```
```
>>> yoda('Hallo') == 'Hallo'
True
```
```
>>> yoda('Hallo mein Name ist Hase') == 'Hase ist Name mein Hallo'
True
```


Tests for q4 (points: None)
```
>>> mask([]) == []
True
```
```
>>> mask([1,2,3,4]) == ['ungerade', 'gerade', 'ungerade', 'gerade']
True
```


Tests for q5 (points: None)
```
>>> median([1,2,3]) == 2
True
```
```
>>> median([1,2,3,5]) == (2+3)/2.0
True
```
```
>>> import random
>>> temp_list1 = [1,2,3,5]
>>> random.shuffle(temp_list1)
>>> median(temp_list1) == (2+3)/2.0
True
```
```
>>> import random
>>> temp_list2 = [-3,1,2,3,5]
>>> random.shuffle(temp_list2)
>>> median(temp_list2) == 2
True
```
